```python

# Initialize Otter
import otter
grader = otter.Notebook("DSA2025Test.ipynb")
```

```python
pip install otter-grader
```

Output:
```
Collecting otter-grader
  Downloading otter_grader-6.1.1-py3-none-any.whl.metadata (3.9 kB)
Collecting click<9.0.0,>=8.1.7 (from otter-grader)
  Downloading click-8.1.8-py3-none-any.whl.metadata (2.3 kB)
Collecting dill>=0.3.0 (from otter-grader)
  Downloading dill-0.3.9-py3-none-any.whl.metadata (10 kB)
Collecting fica>=0.4.1 (from otter-grader)
  Downloading fica-0.4.1-py3-none-any.whl.metadata (2.0 kB)
Collecting ipylab<2.0.0,>=1.0.0 (from otter-grader)
  Downloading ipylab-1.0.0-py3-none-any.whl.metadata (6.7 kB)
Requirement already satisfied: ipython in c:\users\roche\appdata\roaming\python\python313\site-packages (from otter-grader) (9.0.2)
Requirement already satisfied: ipywidgets<9.0.0,>=8.1.5 in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (from otter-grader) (8.1.5)
Requirement already satisfied: jinja2<4.0,>=3.1 in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (from otter-grader) (3.1.6)
Collecting jupytext<2.0.0,>=1.16.
```

```python
pip show otter-grader
```

Output:
```
Name: otter-grader
Version: 6.1.1
Summary: A Python and R autograding solution
Home-page: https://otter-grader.readthedocs.io/
Author: Christopher Pyles
Author-email: cpyles@berkeley.edu
License: BSD-3-Clause
Location: c:\Users\Roche\AppData\Local\Programs\Python\Python313\Lib\site-packages
Requires: click, dill, fica, ipylab, ipython, ipywidgets, jinja2, jupytext, nbconvert, nbformat, pandas, python-on-whales, pyyaml, requests, wrapt
Required-by: 
Note: you may need to restart the kernel to use updated packages.
```

# DSA 2025 Summer School Admittance Check

Thanks for your interest in attending DSA 2025 Ibadan, Nigeria. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Ibadan!

## Instructions
1. Complete each function according to the provided specifications
2. Run all cells to verify your solutions
3. All tests must pass to generate a submission
4. Save your work before submitting

```python
# Run these once ... To be on a safe side
%pip install otter-grader
%pip install numpy
%pip install pandas
%pip install IPython

# Import the good stuff
import IPython
from IPython import get_ipython
import pandas as pd
import numpy as np
import math
import unittest
import otter
from collections import Counter

grader = otter.Notebook()

# Example usage of unittest
class TestMath(unittest.TestCase):
    def test_addition(self):
        self.assertEqual(1 + 1, 2)

# Run the tests
unittest.main(argv=[''], verbosity=2, exit=False)
```

Output:
```
Requirement already satisfied: otter-grader in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (6.1.1)
Requirement already satisfied: click<9.0.0,>=8.1.7 in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (from otter-grader) (8.1.8)
Requirement already satisfied: dill>=0.3.0 in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (from otter-grader) (0.3.9)
Requirement already satisfied: fica>=0.4.1 in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (from otter-grader) (0.4.1)
Requirement already satisfied: ipylab<2.0.0,>=1.0.0 in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (from otter-grader) (1.0.0)
Requirement already satisfied: ipython in c:\users\roche\appdata\roaming\python\python313\site-packages (from otter-grader) (9.0.2)
Requirement already satisfied: ipywidgets<9.0.0,>=8.1.5 in c:\users\roche\appdata\local\programs\python\python313\lib\site-packages (fr
```

<!-- BEGIN QUESTION -->

**Question 1:** 
Write a Python function to return a tuple of even and odd numbers 

1.   List item
2.   List item

given an integer input.</br>
Example: For input 10, the output should be ([2, 4, 6, 8, 10], [1, 3, 5, 7, 9]).

```python
def even_odd_numbers(n):
    """
    Returns a tuple of even and odd numbers up to n.

    Args:
        n (int): The upper limit.

    Returns:
        tuple: A tuple of two lists: (evens, odds).
    """
    evens = [i for i in range(1, n+1) if i % 2 == 0]
    odds = [i for i in range(1, n+1) if i % 2 != 0]
    return (evens, odds)  
print(even_odd_numbers(10))  
```

Output:
```
([2, 4, 6, 8, 10], [1, 3, 5, 7, 9])
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2: Create a dictionary to store the frequency of each word in the paragraph provided.**</br>
Example: For the paragraph "Data science is fun. Data science is useful.", the output should be {'Data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}.

```python
from collections import Counter
import re

def word_frequency(paragraph):
    """
    Returns a dictionary of word frequencies in a paragraph.

    Args:
        paragraph (str): The input paragraph.

    Returns:
        dict: A dictionary with words as keys and their frequencies as values.
    """
    words = re.findall(r'\b\w+\b', paragraph.lower())  # Extract words using regex
    word_counts = Counter(words)  # Count occurrences of each word
    return dict(word_counts)

# Example usage:
paragraph = "Data science is fun. Data science is useful."
print(word_frequency(paragraph))  

```

Output:
```
{'data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3: Extract the values from the dictionary above into a list and sort them in descending order.**</br>
Example: For the dictionary {'a': 3, 'b': 1, 'c': 2}, the output should be [3, 2, 1]

```python
def extract_and_sort_values(dictionary):
    """
    Extracts and sorts dictionary values in descending order.

    Args:
        dictionary (dict): The input dictionary.

    Returns:
        list: A list of sorted values in descending order.
    """
    return sorted(dictionary.values(), reverse=True)

data = {'a': 3, 'b': 1, 'c': 2}
print(extract_and_sort_values(data))

```

Output:
```
[3, 2, 1]
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 4: Merge the two lists (from Questions 3 and 4) into a dictionary where the keys are the sorted keys and the values are the sorted values.**</br>
Example: For lists ['a', 'b', 'c'] and [3, 2, 1], the output should be {'a': 3, 'b': 2, 'c': 1}

```python
def merge_lists_to_dict(keys, values):
    """
    Merges two lists into a dictionary, ensuring that corresponding keys and values match.

    Args:
        keys (list): List of keys.
        values (list): List of values.

    Returns:
        dict: A dictionary with keys and values paired.
    """
    if len(keys) != len(values):
        raise ValueError("Keys and values lists must be of the same length.")

    return dict(zip(keys, values))

# Test case
assert merge_lists_to_dict(['a', 'b', 'c'], [3, 2, 1]) == {'a': 3, 'b': 2, 'c': 1}  # Expected output: {'a': 3, 'b': 2, 'c': 1}

print("Test case passed!")  # This prints only if there's no error.

```

Output:
```
Test case passed!
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5:Write a function least_common_multiple that takes two inputs a and b and returns the least common multiple of the two numbers.**</br>
Example: For input (4, 6), the output should be 12

```python
import math

def least_common_multiple(a, b):
    """
    Returns the least common multiple (LCM) of two numbers.

    Args:
        a (int): First number.
        b (int): Second number.

    Returns:
        int: The LCM of a and b.
    """
    return abs(a * b) // math.gcd(a, b)  # Using GCD to calculate LCM

# Test case
assert least_common_multiple(4, 6) == 12  
assert least_common_multiple(7, 3) == 21  
assert least_common_multiple(10, 15) == 30 



```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6:Write a function get_nearest_farthest that takes in a point of interest (x, y) and a list of points [(x1, y1), (x2, y2), ...] and returns the indices of the nearest and farthest points from the point of interest.** </br>
Example: For pt = (0, 0) and points = [(1, 1), (3, 3), (-1, -1)], the output should be (0, 1)

```python
def get_nearest_farthest(pt, points):
    """
    Returns the indices of the nearest and farthest points from a point of interest.

    Args:
        pt (tuple): The point of interest (x, y).
        points (list): List of points [(x1, y1), (x2, y2), ...].

    Returns:
        tuple: Indices of the nearest and farthest points.
    """
    distances = [( (x - pt[0])**2 + (y - pt[1])**2 ) for x, y in points]
    
    nearest_index = distances.index(min(distances))  # Index of minimum distance
    farthest_index = distances.index(max(distances))  # Index of maximum distance
    
    return (nearest_index, farthest_index)

# Test case
pt = (0, 0)
points = [(1, 1), (3, 3), (-1, -1)]
assert get_nearest_farthest(pt, points) == (0, 1)  #  Should pass

print("All test cases passed!")  # This prints only if everything works correctly.

```

Output:
```
All test cases passed!
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7:Write a function filter_divisible to return a list of numbers between 0 and a number N that are not perfectly divisible by q.** </br>

> Hint: If N is negative, use N = 20. 

Example: For N = 10 and q = 2, the output should be [1, 3, 5, 7, 9]

```python
def filter_divisible(N, q):
    """
    Returns a list of numbers between 0 and N that are not divisible by q.

    Args:
        N (int): The upper limit.
        q (int): The divisor.

    Returns:
        list: A list of non-divisible numbers.
    """
    # If N is negative, set it to 20
    if N < 0:
        N = 20

    # Generate numbers from 0 to N and filter out those divisible by q
    return [num for num in range(N) if num % q != 0]

# Test cases
assert filter_divisible(10, 2) == [1, 3, 5, 7, 9]  #  Should pass
assert filter_divisible(-5, 3) == [1, 2, 4, 5, 7, 8, 10, 11, 13, 14, 16, 17, 19]  #  N = 20
assert filter_divisible(15, 5) == [1, 2, 3, 4, 6, 7, 8, 9, 11, 12, 13, 14]  #  Should pass

print("All test cases passed!")  # This prints only if everything works correctly.

```

Output:
```
All test cases passed!
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8: Write a function flatten_and_unique that takes in a list of lists and outputs a sorted list of unique elements from all sublists.**

```python
def flatten_and_unique(list_of_lists):
    """
    Flattens a list of lists and returns a sorted list of unique elements.

    Args:
        list_of_lists (list): A list of lists.

    Returns:
        list: A sorted list of unique elements.
    """
    # Flatten the list using list comprehension and convert it to a set to get unique values
    unique_elements = set(item for sublist in list_of_lists for item in sublist)

    # Return the sorted list of unique elements
    return sorted(unique_elements)

# Test cases
assert flatten_and_unique([[1, 2, 3], [3, 4, 5], [6, 1]]) == [1, 2, 3, 4, 5, 6]  #  Should pass
assert flatten_and_unique([[10, 20], [20, 30, 40], [10, 50]]) == [10, 20, 30, 40, 50]  #  Should pass
assert flatten_and_unique([[7, 8, 9], [9, 8, 7], [10, 11]]) == [7, 8, 9, 10, 11]  #  Should pass

print("All test cases passed!")  # This prints only if everything works correctly.

```

Output:
```
All test cases passed!
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

<!-- END QUESTION -->

**The Extra Mile!**


Download the dataset Nigeria Food Prices (9.9M) 

> https://data.humdata.org/dataset/wfp-food-prices-for-nigeria





<!-- BEGIN QUESTION -->

**Question 9: Create a new column date_new from the date column, converting it to a datetime format.**

```python
import pandas as pd

# Load the dataset
file_path = r'C:\Users\Roche\Downloads\wfp_food_prices_nga.csv'  # Replace with the actual file path
df = pd.read_csv(file_path,skiprows=[1])
df = df.drop(index=2)
df = df.reset_index(drop=True)

# Display the first few rows of the dataset
print("First 5 rows of the dataset:")
print(df.head())


```

Output:
```
First 5 rows of the dataset:
         date   admin1 admin2        market  latitude  longitude  \
0  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
1  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
2  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
3  2002-01-15  Katsina  Jibia   Jibia (CBM)    13.080      7.240   
4  2002-01-15   Sokoto   Gada  Illela (CBM)    13.645      5.278   

             category      commodity unit priceflag  pricetype currency  \
0  cereals and tubers          Maize   KG    actual  Wholesale      NGN   
1  cereals and tubers         Millet   KG    actual  Wholesale      NGN   
2  cereals and tubers        Sorghum   KG    actual  Wholesale      NGN   
3     pulses and nuts  Beans (niebe)   KG    actual  Wholesale      NGN   
4  cereals and tubers          Maize   KG    actual  Wholesale      NGN   

    price  usdprice  
0  175.92    1.5398  
1  150.18    1.3145  
2  155.61    1.3620  
3  196.87    1.7231  
4
```

```python
import pandas as pd

def create_date_new(df):
    """
    Creates a new column `date_new` from the `date` column, converting it to a datetime format.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.DataFrame: The dataframe with the new `date_new` column.
    """
    df['date_new'] = pd.to_datetime(df['date'], errors='coerce')  # Converts to datetime, handling errors
    return df

```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10:Split the dataframe into two separate dataframes based on the category column: one for cereals and tubers and another for pulses and nuts**

```python
import pandas as pd

def split_dataframes(df):
    """
    Splits the dataframe into two separate dataframes based on the `category` column.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        tuple: A tuple of two dataframes: (cereals_df, pulses_df).
    """
    # Ensure category column exists
    if 'category' not in df.columns:
        raise KeyError("The dataframe does not contain a 'category' column.")

    # Normalize category values (strip spaces and lowercase)
    df['category'] = df['category'].str.strip().str.lower()

    # Print unique categories to check if they match expected values
    print("Unique categories in dataframe:", df['category'].unique())

    # Corrected filtering using lowercase
    cereals_df = df[df['category'] == 'cereals and tubers'].copy()
    pulses_df = df[df['category'] == 'pulses and nuts'].copy()

    # Ensure the function always returns a valid DataFrame
    if cereals_df.empty:
        print("Warning: 'Cereals and Tubers' dataframe is empty.")
    if pulses_df.empty:
        print("Warning: 'Pulses and Nuts' dataframe is empty.")

    return cereals_df, pulses_df  # Ensure both are always returned

```

```python
grader.check("q10")
```

Output:
```
q10 results:
    q10 - 1 result:
         Test case failed
        Trying:
            (cereals_df, pulses_df) = split_dataframes(df)
        Expecting nothing
        **********************************************************************
        Line 1, in q10 0
        Failed example:
            (cereals_df, pulses_df) = split_dataframes(df)
        Expected nothing
        Got:
            Unique categories in dataframe: ['cereals and tubers' 'pulses and nuts' 'non-food' 'oil and fats'
             'meat, fish and eggs' 'milk and dairy' 'vegetables and fruits'
             'miscellaneous food']
        Trying:
            assert len(cereals_df) > 0, 'Cereals and tubers dataframe is empty'
        Expecting nothing
        ok

    q10 - 2 result:
         Test case failed
        Trying:
            assert len(pulses_df) > 0, 'Pulses and nuts dataframe is empty'
        Expecting nothing
        **********************************************************************
        Line 1, i
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 11:Calculate the mean, median, and mode of the price and usdprice columns for each category**

```python
import pandas as pd

def calculate_price_stats(df):
    """
    Calculates the mean, median, and mode of the `price` and `usdprice` columns for each `category`.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        tuple: A tuple of two dataframes: (price_stats, usdprice_stats).
    """
    # Ensure required columns exist
    required_columns = {'category', 'price', 'usdprice'}
    if not required_columns.issubset(df.columns):
        raise KeyError(f"Missing required columns: {required_columns - set(df.columns)}")
    
    # Group by category and compute statistics
    price_stats = df.groupby('category')['price'].agg(['mean', 'median'])
    usdprice_stats = df.groupby('category')['usdprice'].agg(['mean', 'median'])

    # Compute mode separately since it may have multiple values
    price_mode = df.groupby('category')['price'].apply(lambda x: x.mode().iloc[0] if not x.mode().empty else None)
    usdprice_mode = df.groupby('category')['usdprice'].apply(lambda x: x.mode().iloc[0] if not x.mode().empty else None)

    # Add mode to the stats DataFrames
    price_stats['mode'] = price_mode
    usdprice_stats['mode'] = usdprice_mode

    return price_stats, usdprice_stats

```

```python
grader.check("q11")
```

Output:
```
q11 results:
    q11 - 1 result:
         Test case passed

    q11 - 2 result:
         Test case failed
        Trying:
            assert not usdprice_stats.empty, 'USD price statistics dataframe is empty'
        Expecting nothing
        **********************************************************************
        Line 1, in q11 1
        Failed example:
            assert not usdprice_stats.empty, 'USD price statistics dataframe is empty'
        Exception raised:
            Traceback (most recent call last):
              File "c:\Users\Roche\AppData\Local\Programs\Python\Python313\Lib\doctest.py", line 1395, in __run
                exec(compile(example.source, filename, "single",
                ~~~~^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
                             compileflags, True), test.globs)
                             ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
              File "<doctest q11 1[0]>", line 1, in <module>
                assert not usdprice_stats.empty, '
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 12:Merge the two dataframes (from Question 11) back into one dataframe and save it as `merged_data.csv`.**

```python
def merge_dataframes(cereals_df, pulses_df):
    """
    Merges the two dataframes back into one dataframe and saves it as `merged_data.csv`.

    Args:
        cereals_df (pd.DataFrame): The cereals and tubers dataframe.
        pulses_df (pd.DataFrame): The pulses and nuts dataframe.

    Returns:
        pd.DataFrame: The merged dataframe.
    """

```

```python
grader.check("q12")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 13:Open the `merged_data.csv` file and select only the `date_new`, `market`, `commodity`, `price`, and `usdprice` columns**

```python
def select_columns():
    """
    Opens the `merged_data.csv` file and selects only the `date_new`, `market`, `commodity`, `price`, and `usdprice` columns.

    Returns:
        pd.DataFrame: The dataframe with selected columns.
    """
```

```python
grader.check("q13")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 14:Group the data by `admin1` (state) and calculate the average `price` and `usdprice` for each state**

```python
def calculate_state_avg_prices(df):
    """
    Groups the data by `admin1` (state) and calculates the average `price` and `usdprice` for each state.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.DataFrame: The dataframe with average prices by state.
    """
```

```python
grader.check("q14")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 15:Identify the top 5 markets with the highest average `usdprice` for `Rice (imported)`.**

```python
def identify_top_markets(df):
    """
    Identifies the top 5 markets with the highest average `usdprice` for `Rice (imported)`.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.Series: A series with the top 5 markets and their average `usdprice`.
    """ 
```

```python
grader.check("q15")
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Complete each function according to the provided specifications. Make sure all tests pass.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert even_odd_numbers(10) == ([2, 4, 6, 8, 10], [1, 3, 5, 7, 9])
```
```
>>> assert even_odd_numbers(5) == ([2, 4], [1, 3, 5])
```


Tests for q10 (points: 1)
```
>>> (cereals_df, pulses_df) = split_dataframes(df)
>>> assert len(cereals_df) > 0, 'Cereals and tubers dataframe is empty'
```
```
>>> assert len(pulses_df) > 0, 'Pulses and nuts dataframe is empty'
```


Tests for q11 (points: 1)
```
>>> (price_stats, usdprice_stats) = calculate_price_stats(df)
>>> assert not price_stats.empty, 'Price statistics dataframe is empty'
```
```
>>> assert not usdprice_stats.empty, 'USD price statistics dataframe is empty'
```


Tests for q12 (points: 1)
```
>>> merged_df = merge_dataframes(cereals_df, pulses_df)
>>> assert not merged_df.empty, 'Merged dataframe is empty'
```
```
>>> assert len(merged_df) == len(cereals_df) + len(pulses_df), 'Merged dataframe has incorrect number of rows'
```


Tests for q13 (points: 1)
```
>>> selected_columns = select_columns()
>>> assert not selected_columns.empty, 'Selected columns dataframe is empty'
```
```
>>> assert list(selected_columns.columns) == ['date_new', 'market', 'commodity', 'price', 'usdprice'], 'Incorrect columns selected'
```


Tests for q14 (points: 1)
```
>>> state_avg_prices = calculate_state_avg_prices(df)
>>> assert not state_avg_prices.empty, 'State average prices dataframe is empty'
```
```
>>> assert 'price' in state_avg_prices.columns, 'Price column not found in state average prices'
```


Tests for q15 (points: 1)
```
>>> top_markets = identify_top_markets(df)
>>> assert not top_markets.empty, 'Top markets dataframe is empty'
```
```
>>> assert len(top_markets) == 5, 'Incorrect number of top markets'
```


Tests for q2 (points: 1)
```
>>> assert word_frequency('Data science is fun. Data science is useful.') == {'data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}
```
```
>>> assert word_frequency('Hello hello world') == {'hello': 2, 'world': 1}
```


Tests for q3 (points: 1)
```
>>> assert extract_and_sort_values({'a': 3, 'b': 1, 'c': 2}) == [3, 2, 1]
```
```
>>> assert extract_and_sort_values({'x': 10, 'y': 5, 'z': 20}) == [20, 10, 5]
```


Tests for q4 (points: 1)
```
>>> assert merge_lists_to_dict(['a', 'b', 'c'], [3, 2, 1]) == {'a': 3, 'b': 2, 'c': 1}
```
```
>>> assert merge_lists_to_dict(['x', 'y'], [10, 20]) == {'x': 10, 'y': 20}
```


Tests for q5 (points: 1)
```
>>> assert least_common_multiple(4, 6) == 12
```
```
>>> assert least_common_multiple(5, 10) == 10
```


Tests for q6 (points: 1)
```
>>> assert get_nearest_farthest((0, 0), [(1, 1), (3, 3), (-1, -1)]) == (0, 1)
```
```
>>> assert get_nearest_farthest((0, 0), [(2, 2), (1, 1), (3, 3)]) == (1, 2)
```


Tests for q7 (points: 1)
```
>>> assert filter_divisible(10, 2) == [1, 3, 5, 7, 9]
```
```
>>> assert filter_divisible(15, 3) == [1, 2, 4, 5, 7, 8, 10, 11, 13, 14]
```


Tests for q8 (points: 1)
```
>>> assert flatten_and_unique([[2, 13, 44], [6, 7, 2]]) == [2, 6, 7, 13, 44]
```
```
>>> assert flatten_and_unique([[1, 2, 3], [3, 4, 5]]) == [1, 2, 3, 4, 5]
```


Tests for q9 (points: 1)
```
>>> df = create_date_new(df)
>>> assert 'date_new' in df.columns, "Column 'date_new' not found in the dataframe"
```
```
>>> assert df['date_new'].dtype == 'datetime64[ns]', "Column 'date_new' is not in datetime format"
```
